---
title: Install Docker
description: Testing the install of docker

---
## Sample exercise

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: ecaded3183
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}
pip install docker
```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
